# QBUS1040: Foundations of Business Analytics
## Homework 3
Semester 2, 2019

This homework consists of eight problems that require you to submit a written response and a coding component. The coding component is described in this file. Please answer all questions in the space provided in this homework file. When you are done, save the file and upload it to Canvas. You should submit one file for the whole homework.  

__Good luck!__

# Problem 4(a)

The variable `vec_list` contains 15 vectors. Run the Gram-Schmidt algorithm on the given set of vectors in the variable `vec_list` and return the variable `q_list` as the vectors $q_1, \ldots, q_{15}$.

<!-- grade_id: cell-a6440e82f4f07309, points: 0, grade: false, solution: true -->
```python
import numpy as np
np.random.seed(seed=100)
A = np.random.random((15, 15))
vec_list = []
for item in A:
    vec_list.append(item)
q_list = []

#Gram-Schmidt Algorithm    
import math

def scalar_division(v, n):
    i = 0
    new_v = []
    while i < len(v):
        x = v[i] / n
        new_v.append(x)
        i += 1
    return new_v

def scalar_mult(v, n):
    i = 0
    new_v = []
    while i < len(v):
        x = v[i]*n
        new_v.append(x)
        i +=1
    return new_v

def vector_alg(v1, v2, flag):
    i = 0
    new_v = []
    if flag == "plus":
        while i < len(v1):
            x = v1[i] + v2[i]
            new_v.append(x)
            i +=1
    elif flag == "minus":
        while i < len(v1):
            x = v1[i] - v2[i]
            new_v.append(x)
            i += 1
    return new_v

def norm(vector_x):
    sum_squared = 0
    i = 0
    while i < len(vector_x):
        sum_squared += vector_x[i]**2
        i += 1
    norm = math.sqrt(sum_squared)
    return norm 

def inner(v1, v2):
    i = 0
    total = 0
    while i < len(v1):
        total += v1[i]*v2[i]
        i +=1 
    return total

def gs_alg(a_list):
    qi_list = []
    i = 0
    while i < len(a_list):
        #Orthogonalisation
        qi_tilde = orthogornalisation(a_list[i], qi_list)
        
        #Test for independence 
        if norm(qi_tilde) < 1e-6:
            return qi_list
        
        #Normalisation
        else:
            qi = scalar_division(qi_tilde, norm(qi_tilde))
            qi_list.append(qi)
        i += 1
    return qi_list
        
def orthogornalisation(a, qi_list):
    qi_tilde = a
    if len(qi_list) == 0:
        return qi_tilde
    else:
        k = 0
        while k < len(qi_list):
            qi_tilde = vector_alg(qi_tilde, scalar_mult(qi_list[k], inner(qi_list[k], a)), "minus")
            k += 1
    return qi_tilde

q_list = np.array(gs_alg(vec_list))
#print(q_list)
```

Output:
```
[[ 2.82679345e-01  1.44807802e-01  2.20834124e-01  4.39452690e-01
   2.45474981e-03  6.32402777e-02  3.48923790e-01  4.29608746e-01
   7.11147917e-02  2.99163650e-01  4.63665835e-01  1.08826980e-01
   9.64077719e-02  5.63776772e-02  1.14286674e-01]
 [ 3.46432736e-01  4.08960951e-01 -1.71351028e-01  1.26138307e-02
   2.01095344e-01  2.36202804e-01  2.27774777e-01  2.70313720e-02
   1.54219563e-01 -2.75030225e-01 -3.50952238e-01 -1.43398893e-01
   5.74861876e-02  5.16981619e-01 -1.43670455e-01]
 [-2.40489258e-02  1.04498954e-01 -1.05657627e-01 -2.18980336e-01
  -4.87044072e-02  4.42894903e-01  2.01264981e-01 -4.27986181e-01
   4.63273716e-01  1.61245372e-01  1.22552100e-01  3.53779543e-01
   2.64529434e-01 -2.27482214e-01  6.82081442e-02]
 [-1.31410502e-01  1.07527832e-01  1.93425534e-01 -1.30327166e-01
   7.08063631e-01  2.67561419e-01  1.44482759e-02  1.33725954e-01
  -6.21150121e-02  3.77222919e-02 -9.98016641e-02 -1.68233294e-01
  -1.88208338e-01 -2.81970825e-01  4.08976591e-01]
 [-1
```

<!-- grade_id: cell-657dd5c7268f3183, points: 2.5, grade: true, solution: false -->
```python
# Assertion cell (Q4a) - please do not delete or move this cell
```

# Problem 4(b)

Run the Gram-Schmidt algorithm on vectors $q_1,\ldots, q_{15}$ obtained from the previous part and return the variable `z_list` as the vectors $z_1,\ldots,z_{15}$ obtained from running the Gram-Schmidt for the second time.

<!-- grade_id: cell-a6440e82f4f07308, points: 0, grade: false, solution: true -->
```python
#Gram-Schmidt Algorithm    
z_list = []
# use q_list from cell above as input

z_list = np.array(gs_alg(q_list))
#print(z_list)
```

Output:
```
[[ 2.82679345e-01  1.44807802e-01  2.20834124e-01  4.39452690e-01
   2.45474981e-03  6.32402777e-02  3.48923790e-01  4.29608746e-01
   7.11147917e-02  2.99163650e-01  4.63665835e-01  1.08826980e-01
   9.64077719e-02  5.63776772e-02  1.14286674e-01]
 [ 3.46432736e-01  4.08960951e-01 -1.71351028e-01  1.26138307e-02
   2.01095344e-01  2.36202804e-01  2.27774777e-01  2.70313720e-02
   1.54219563e-01 -2.75030225e-01 -3.50952238e-01 -1.43398893e-01
   5.74861876e-02  5.16981619e-01 -1.43670455e-01]
 [-2.40489258e-02  1.04498954e-01 -1.05657627e-01 -2.18980336e-01
  -4.87044072e-02  4.42894903e-01  2.01264981e-01 -4.27986181e-01
   4.63273716e-01  1.61245372e-01  1.22552100e-01  3.53779543e-01
   2.64529434e-01 -2.27482214e-01  6.82081442e-02]
 [-1.31410502e-01  1.07527832e-01  1.93425534e-01 -1.30327166e-01
   7.08063631e-01  2.67561419e-01  1.44482759e-02  1.33725954e-01
  -6.21150121e-02  3.77222919e-02 -9.98016641e-02 -1.68233294e-01
  -1.88208338e-01 -2.81970825e-01  4.08976591e-01]
 [-1
```

<!-- grade_id: cell-657dd5c7268f3182, points: 2.5, grade: true, solution: false -->
```python
# Assertion cell (Q4b) - please do not delete or move this cell
```

# Problem 4(c)

Compute the distances between vectors $q_1$ and $z_1$, $q_2$ and $z_2$, ..., $q_{15}$ and $z_{15}$. Assign your answer (a 15-vector) as a numpy array to the variable `ans_1c`.

<!-- grade_id: cell-ce31e9e8203fa7e7, points: 0, grade: false, solution: true -->
```python
def distance(v1, v2):
    diff = vector_alg(v1, v2, "minus")
    norm_diff = norm(diff)
    return norm_diff

ans_ls = []
i = 0
while i < len(q_list):
    x = distance(q_list[i], z_list[i])
    ans_ls.append(x)
    i += 1

ans_4c = np.array(ans_ls)
print(ans_4c)
```

Output:
```
[0.00000000e+00 2.11045349e-16 1.68866220e-16 5.06447164e-16
 8.71530989e-16 7.78866814e-16 1.66125293e-15 4.60519391e-15
 6.99093142e-15 2.79002216e-15 4.25221226e-15 3.52360045e-15
 2.57206789e-14 2.94397617e-14 5.17039378e-14]
```

<!-- grade_id: cell-2f5d7886ea0ecb0c, points: 2.5, grade: true, solution: false -->
```python
# Assertion cell (Q4c) - please do not delete or move this cell
```

# Problem 8

_Orthogonality_

Suppose the $n$-vectors $a$, $b$, and $c$ satisfy $a \perp c$ and $b \perp c$. Which of the following statements must hold? (That is, are true for any $a$, $b$, and $c$ that satisfy $a \perp c$, $b \perp c$.) Assign __True__ or __False__ to the variable provided in the answer cell.

(a) $a \perp b$.

<!-- grade_id: cell-69e537713d587d2a, points: 0, grade: false, solution: true -->
```python
#Answer
ans_8a = False

```

<!-- grade_id: cell-75c83b9550c02c83, points: 3, grade: true, solution: false -->
```python
# Assertion cell (Q8a) - please do not delete or move this cell
```

(b) $(a+b) \perp c$.

<!-- grade_id: cell-feb94671d69aa1ae, points: 0, grade: false, solution: true -->
```python
#Answer
ans_8b = True

```

<!-- grade_id: cell-63087bc7b3037b9c, points: 3, grade: true, solution: false -->
```python
# Assertion cell (Q8b) - please do not delete or move this cell
```

(c) $(a+c) \perp (b+c)$.

<!-- grade_id: cell-433246a5b5c70044, points: 0, grade: false, solution: true -->
```python
#Answer
ans_8c = False

```

<!-- grade_id: cell-7aee7c773159ad58, points: 4, grade: true, solution: false -->
```python
# Assertion cell (Q8c) - please do not delete or move this cell
```

<!-- grade_id: cell-71ff7fead64b1421, points: 0, grade: true, solution: false -->
```python
# Adhoc assertion cell - please do not delete or move this cell
```

<!-- grade_id: cell-80dff0dc5b005b81, points: 0, grade: true, solution: false -->
```python
# Adhoc assertion cell - please do not delete or move this cell
```