# QBUS1040: Foundations of Business Analytics
## Homework 2
Semester 2, 2019

Please answer all questions in the space provided in this homework file. When you are done, save the file and upload it to Canvas. You should submit one file for the whole homework.  

__Good luck!__

# Problem 6
_Creating vectors in Python_

In each of the parts below, create vector $a$ such that $a^T x$ provides the correct result according to the description, for the given vector $x$ in each part. Assign your vector $a$ as a numpy array to the variable provided in the answer cell of each part.

(a) $a^T x$ extracts (is equal to) the 7th entry of the $10$-vector $x$

<!-- grade_id: cell-5b097410f9406910, points: 0, grade: false, solution: true -->
```python
import numpy as np
# Answer (Q6a)

# YOUR CODE HERE

ans_6a = np.array([0,0,0,0,0,0,1,0,0,0])
```

<!-- grade_id: cell-71461519b3d9cb64, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q6a) - please do not delete or move this cell
from nose.tools import assert_equal
import numpy as np
rand_vec_x = np.random.rand(10) #I create a random vector x
assert_equal(np.matmul(ans_6a,rand_vec_x),rand_vec_x[6]) # check if the inner product equals the seventh element
```

(b) $a^T x$ is the weighted average of a $3$-vector $x$, assigning weights $0.3$ to the first component, $0.4$ to the second, and $0.3$ to the third.

<!-- grade_id: cell-867f8b2bb748c38c, points: 0, grade: false, solution: true -->
```python
# Answer (Q6b)
ans_6b = np.array([0.3, 0.4, 0.3])
# YOUR CODE HERE

```

<!-- grade_id: cell-c49ef136e9a76e02, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q6b) - please do not delete or move this cell
```

(c) $a^T x$ (with $x$ a $22$-vector) is the sum of $x_i$ for $i$ a multiple of $4$, minus the sum of $x_i$ for $i$ a multiple of $7$.


<!-- grade_id: cell-d293d0ff64f778b7, points: 0, grade: false, solution: true -->
```python
# Answer (Q6c)
ans_6c = np.array([0, 0, 0, 1, 0, 0, -1, 1, 0, 0, 0, 1, 0, -1, 0, 1, 0, 0, 0, 1, -1, 0])
# YOUR CODE HERE

```

<!-- grade_id: cell-10715472e9c3fe96, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q6c) - please do not delete or move this cell
```

(d) $a^T x$ (with $x$ an $11$-vector) is the average of the middle three entries of $x$, i.e., entries $5$ to $7$.

<!-- grade_id: cell-7733ee1dc195bb4f, points: 0, grade: false, solution: true -->
```python
# Answer (Q6d)
ans_6d = np.array([0, 0, 0, 0, 1/3, 1/3, 1/3, 0, 0, 0, 0])
# YOUR CODE HERE

```

<!-- grade_id: cell-bfe6d782d0792d6d, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q6d) - please do not delete or move this cell
```

# Problem 7
_Linear or not?_

Determine whether each of the following scalar-valued functions of $n$-vectors is linear. 

Assign __Boolean value__ `True` or `False` to the variable provided in the answer cell. Answer  `True` if the statement is linear and `False` otherwise.

(a) The spread of values of the vector, defined as $f(x) = max_k x_k − min_k x_k$.


<!-- grade_id: cell-8307552fb6b867fe, points: 0, grade: false, solution: true -->
```python
# Answer (Q7a)
ans_7a = False
# YOUR CODE HERE

import random
from random import randint

#creating 2 random 4-vectors 
ls = []
i = 0 
random.seed(randint(0,99))
while i <= 1:
    vect = [randint(0,99), randint(0,99), randint(0,99), randint(0,99)]
    ls.append(vect)
    i += 1
print(ls)

x = ls[0]
y = ls[1]

a = randint(0,99)
b = randint(0,99)

def spread(ls):
    return max(ls) - min(ls)

######################
ax = []
i = 0
while i < len(x):
    n = x[i]*a
    ax.append(n)
    i += 1
    
by = []
i = 0
while i < len(y):
    n = y[i]*b
    by.append(n)
    i += 1
    
ax_by = []
i = 0
while i < len(ax):
    n = ax[i] + by[i]
    ax_by.append(n)
    i +=1

result1 = spread(ax_by)
print(result1)
######################

######################
f_x = spread(x) #scalar
f_y = spread(y) #scalar

afx = a*f_x
bfy = b*f_y
result2 = afx + bfy
print(result2)
######################    
```

<!-- grade_id: cell-a2daec160aaa4f02, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q7a) - please do not delete or move this cell
```

(b) The difference of the first element and the last, $f(x) = x_1 − x_n$.

<!-- grade_id: cell-b421b88e11fe2217, points: 0, grade: false, solution: true -->
```python
# Answer (Q7b)
ans_7b = True 

# YOUR CODE HERE
import random
from random import randint

#creating 2 random 4-vectors 
ls = []
i = 0 
random.seed(randint(0,99))
while i <= 1:
    vect = [randint(0,99), randint(0,99), randint(0,99), randint(0,99)]
    ls.append(vect)
    i += 1
print(ls)

x = ls[0]
y = ls[1]

a = randint(0,99)
b = randint(0,99)


def diff(ls):
    return ls[0] - ls[-1]

######################
ax = []
i = 0
while i < len(x):
    n = x[i]*a
    ax.append(n)
    i += 1
    
by = []
i = 0
while i < len(y):
    n = y[i]*b
    by.append(n)
    i += 1
    
ax_by = []
i = 0
while i < len(ax):
    n = ax[i] + by[i]
    ax_by.append(n)
    i +=1

result1 = diff(ax_by)
print(result1)
######################

######################
f_x = diff(x) #scalar
f_y = diff(y) #scalar

afx = a*f_x
bfy = b*f_y
result2 = afx + bfy
print(result2)
######################
```

<!-- grade_id: cell-df4532b5b986ec4d, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q7b) - please do not delete or move this cell
```

(c) The median of an $n$-vector, where we will assume $n = 2k + 1$ is odd. The median of
the vector $x$ is defined as the $(k + 1)$st largest number among the entries of $x$. For
example, the median of $(−7.1, 3.2, −1.5)$ is $−1.5$.


<!-- grade_id: cell-b0241254f7b25c1b, points: 0, grade: false, solution: true -->
```python
# Answer (Q7c)
ans_7c = False

# YOUR CODE HERE
#creating 2 random 5-vectors 
import random
from random import randint

ls = []
i = 0 
random.seed(randint(0,99))
while i <= 1:
    vect = [randint(0,99), randint(0,99), randint(0,99), randint(0,99), randint(0,99)]
    ls.append(vect)
    i += 1
print(ls)

x = ls[0]
y = ls[1]

a = randint(0,99)
b = randint(0,99)

def median(vector):
    ls = sorted(vector)
    k = int((len(vector)-1) / 2) 
    index = k + 1 #3rd largest number 
    median = ls[-index] #when sorted, the last element will be the largest => 3rd largest = -index
    return median 

######################
ax = []
i = 0
while i < len(x):
    n = x[i]*a
    ax.append(n)
    i += 1
    
by = []
i = 0
while i < len(y):
    n = y[i]*b
    by.append(n)
    i += 1
    
ax_by = []
i = 0
while i < len(ax):
    n = ax[i] + by[i]
    ax_by.append(n)
    i +=1

result1 = median(ax_by)
print(result1)
######################

######################
f_x = median(x) #scalar
f_y = median(y) #scalar

afx = a*f_x
bfy = b*f_y
result2 = afx + bfy
print(result2)
######################
```

<!-- grade_id: cell-37b87b9c2a6e65b5, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q7c) - please do not delete or move this cell
```

(d) The average of the entries with odd indices, minus the average of the entries with
even indices. You can assume that $n = 2k$ is even.


<!-- grade_id: cell-647de22ba0294e6c, points: 0, grade: false, solution: true -->
```python
# Answer (Q7d)
ans_7d = True 
# YOUR CODE HERE
import random
from random import randint

#creating 2 random 4-vectors 
ls = []
i = 0 
random.seed(randint(0,99))
while i <= 1:
    vect = [randint(0,99), randint(0,99), randint(0,99), randint(0,99)]
    ls.append(vect)
    i += 1
print(ls)

x = ls[0]
y = ls[1]

a = randint(0,99)
b = randint(0,99)

def func(ls):
    if len(ls) == 0:
        return 0
    total_even = 0
    count_even = 0
    total_odd = 0
    count_odd = 0
    i = 0
    while i < len(ls):
        if i % 2 == 0:
            count_even += 1
            total_even += ls[i]
        elif i % 2 != 0:
            count_odd += 1
            total_odd += ls[i]
        i += 1
    avg_even = total_even/count_even
    avg_odd = total_odd/count_odd
    return avg_odd - avg_even

######################
ax = []
i = 0
while i < len(x):
    n = x[i]*a
    ax.append(n)
    i += 1
    
by = []
i = 0
while i < len(y):
    n = y[i]*b
    by.append(n)
    i += 1
    
ax_by = []
i = 0
while i < len(ax):
    n = ax[i] + by[i]
    ax_by.append(n)
    i +=1

result1 = func(ax_by)
print(result1)
######################

######################
f_x = func(x) #scalar
f_y = func(y) #scalar

afx = a*f_x
bfy = b*f_y
result2 = afx + bfy
print(result2)
######################
```

<!-- grade_id: cell-e1822ed7ade8214f, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q7d) - please do not delete or move this cell
```

(e) Vector extrapolation, defined as $x_n + (x_n − x_{n−1})$, for $n \geq 2$. (This is a simple
prediction of what $x_{n+1}$ would be, based on a straight line drawn through $x_n$ and
$x_{n−1}$.)

<!-- grade_id: cell-c2cbfc55c4215081, points: 0, grade: false, solution: true -->
```python
# Answer (Q7e)
ans_7e = False

# YOUR CODE HERE
    

```

<!-- grade_id: cell-5fd447035fe23470, points: 2, grade: true, solution: false -->
```python
# Assertion cell (Q7e) - please do not delete or move this cell
```

<!-- grade_id: cell-3a21b5160e2128b6, points: 0, grade: true, solution: false -->
```python
# Adhoc assertion cell - please do not delete or move this cell
```

<!-- grade_id: cell-26a6bb2597f54d9c, points: 0, grade: true, solution: false -->
```python
# Adhoc assertion cell - please do not delete or move this cell
```

<!-- grade_id: cell-211f17e7b672fc9b, points: 0, grade: true, solution: false -->
```python
# Adhoc assertion cell - please do not delete or move this cell
```